```python
# Initialize Otter
import otter
grader = otter.Notebook("Assignment4_SVM.ipynb")
```

# CS385/CSD3185/CSD3186: Assignment 4 SVM

## Deliverables
Your submission for this assignment should be __ONE__ file - a zip file generated by the `grader.export(_)` function down below under 'Submission'. This function will auto create the zip file which contains this particular completed notebook file for you.  

Subsequently, rename your zip file like this: __coursecode_A4_your_full_name.zip__  
Eg. CS385_A4_john_doe.zip  

To complete this assignment, you should follow instructions in Section Tasks.

## IMPORTANT! READ THIS BEFORE STARTING...
- DO NOT delete existing cells, but you can add more cells in between.
- DO NOT modify the content of the existing cells unless otherwise stated.
- Run the cell with `grader.check(_)` to check your solutions whenever you have completed each tasks.
- Follow the file naming convention for the zip file as spelled out above strictly.
- DO NOT rename this notebook file. It shall be 'assignment1.ipynb'. 

Please adhere strictly to the instructions as stated above as failure to do so might result in deduction of marks by the autograder.
___

# Assignment Overview: Machine Learning Model Comparison

In this assignment, you will apply **Support Vector Machines (SVM)** alongside previously learned machine learning models—**Logistic Regression and Decision Tree**—to perform a comparative analysis. The goal is to evaluate and compare the performance of these models on a given dataset using various evaluation metrics. Through this process, you will gain a deeper understanding of how different classification algorithms perform under different conditions and learn to fine-tune models for optimal results.

## Notebook Structure:
Step 1: Data Loading & Cleaning [_2 marks_]

Step 2: Exploratory Data Analysis (EDA) [_10 marks_]

Step 3: Feature Engineering [_10 marks_]

Step 4: Data Preprocessing & Feature Selection [_64 marks_]

Step 5: Model Training & Selection [_14 marks_]

## Heart Disease Dataset (UCI)

The Heart Disease dataset is a widely used dataset from the UCI Machine Learning Repository. It contains medical data used to predict the presence of heart disease in patients. The dataset consists of multiple patient attributes, including demographic, clinical, and laboratory test results.

### Dataset Details  
- **Number of Features:** 13  
- **Target Variable:** `target` (1 = presence of heart disease, 0 = absence of heart disease)  
- **Type:** Classification  

### Features  
| Feature | Description |
|---------|------------|
| `age` | Age of the patient (years) |
| `sex` | Sex of the patient (1 = male, 0 = female) |
| `cp` | Chest pain type (0-3, different types of angina and pain) |
| `trestbps` | Resting blood pressure (mm Hg) |
| `chol` | Serum cholesterol (mg/dL) |
| `fbs` | Fasting blood sugar > 120 mg/dL (1 = true, 0 = false) |
| `restecg` | Resting electrocardiographic results (0-2) |
| `thalach` | Maximum heart rate achieved |
| `exang` | Exercise-induced angina (1 = yes, 0 = no) |
| `oldpeak` | ST depression induced by exercise relative to rest |
| `slope` | Slope of the peak exercise ST segment (0-2) |
| `ca` | Number of major vessels colored by fluoroscopy (0-3) |
| `thal` | Thalassemia type (1 = normal, 2 = fixed defect, 3 = reversible defect) |
| `target` | Presence of heart disease (1 = disease, 0 = no disease) |

# Import Libraries

```python
import numpy as np
import pandas as pd
import seaborn as sns
import matplotlib.pyplot as plt

import warnings
# Suppress warnings
warnings.filterwarnings("ignore")
```

# Step1: Data Loading & Cleaning
**Tasks:**

- Load the dataset and check for **duplicates**.

- Remove duplicates (**1 duplicate exists** in `heart.csv`).

```python
# Load data
df = pd.read_csv("heart.csv")

# Check duplicates
print("Initial shape:", df.shape[0])
print("Duplicate rows:", df.shape[0] - df.drop_duplicates().shape[0])  

# Drop duplicates
df = df.drop_duplicates()
```

Output:
```
Initial shape: 303
Duplicate rows: 1
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

```python
df.head()
```

Output:
```
age  sex  cp  trestbps  chol  fbs  restecg  thalach  exang  oldpeak  slope  \
0   63    1   3       145   233    1        0      150      0      2.3      0   
1   37    1   2       130   250    0        1      187      0      3.5      0   
2   41    0   1       130   204    0        0      172      0      1.4      2   
3   56    1   1       120   236    0        1      178      0      0.8      2   
4   57    0   0       120   354    0        1      163      1      0.6      2   

   ca  thal  target  
0   0     1       1  
1   0     2       1  
2   0     2       1  
3   0     2       1  
4   0     2       1
```

# Step 2: Exploratory Data Analysis (EDA)
**Tasks:**

- Plot a histogram of **age** with **20 bins**.

- Plot **class distribution** of the target variable (**target**).

- Visualize **correlations** between features and the target using a **heatmap**.

- Identify **outliers** in **Cholesterol (chol)** and **ST depression (oldpeak)** using **boxplots**.

Customize labels and titles for clarity.

```python
import matplotlib.pyplot as plt
import seaborn as sns

# Histogram of age
counts, bin_edges, _ = ...
# Customize labels and titles
...
```

```python
# Class distribution (countplot)
ax_count = ...
# Customize title
...
```

```python
# Correlation heatmap
# Ensure each cell displays the correlation value (annot)
# The values are formatted to one decimal place (fmt)
ax = ...
```

```python
# Outliers (boxplots)
plt.figure(figsize=(10, 4))
plt.subplot(1, 2, 1)
ax1 = ...
plt.subplot(1, 2, 2)
ax2 = ...
```

```python
grader.check("q2")
```

# Step 3: Feature Engineering
**Tasks:**

- **Bin** `age` into groups: `<=45`, `45-55`, `>55`, labelled as **'young'**, **'middle'**, **'senior'**. (Hint: use pd.cut())

- **Create interaction term** `age_chol` (**age × cholesterol**).

```python
# Bin Age into Groups
df['age_group'] = ...

# Create Interaction Term
df['age_chol'] = ...
```

```python
grader.check("q3")
```

# Step 4: Data Preprocessing & Feature Selection
**Tasks:**

- **Split Data** → Features (`X`) & Target (`y`)  
- **Train-Test Split** → 80% train, 20% test, use `random_state = 42` and `stratify = y`
- **Handle Outliers (Post Split!)** → IQR capping 'for `chol` & `oldpeak`  
- **One-Hot Encode** categorical variables (`cp, restecg, slope, thal, age_group`)  
- **Standardize Numerical Features** → `age, trestbps, chol, oldpeak, age_chol`  
- **Feature Selection (RFE + Logistic Regression)** → Keep top 10 features  (codes provided)

```python
from sklearn.model_selection import train_test_split

# Split Data into Features (X) and Target (y)
X = ...
y = ...

# Train-Test Split (80% Train, 20% Test)
# Use random_state = 42 for reproducibility.
# Use stratify = y to ensure the train and test sets have similar class distributions.
X_train, X_test, y_train, y_test = ...
```

```python
# Data Preprocessing: Handling Outliers with IQR Method

def calculate_outlier_bounds(df, columns):
    """
    Calculate IQR-based outlier bounds using ONLY training data to prevent data leakage.
    
    IQR (Interquartile Range) Method:
    - Q1 (25th percentile): The value below which 25% of the observations are found
    - Q3 (75th percentile): The value below which 75% of the observations are found
    - IQR = Q3 - Q1: The range of the middle 50% of the data
    - Lower bound = Q1 - 1.5 * IQR: Values below this are considered outliers
    - Upper bound = Q3 + 1.5 * IQR: Values above this are considered outliers
    
    Parameters:
    -----------
    df : pandas.DataFrame
        The training DataFrame containing numerical columns
    columns : list
        List of column names to calculate IQR bounds for
        
    Returns:
    --------
    bounds : dict
        Dictionary with column names as keys and their corresponding lower/upper bounds
    """
    # Initialize empty dictionary to store bounds for each column
    bounds = {}
   
    
    # Calculate bounds for each specified column
    for col in columns:
        ...
        
        # Store bounds for this column in the dictionary
        bounds[col] = {
            'lower': lower_bound,
            'upper': upper_bound
        }
    
    return bounds


def apply_outlier_bounds(df, columns, bounds):
    """
    Apply pre-calculated outlier bounds to a dataset (training or test).
    
    This function ensures we apply the SAME bounds to both training and test data,
    preventing data leakage by not calculating new bounds from test data.
    
    Parameters:
    -----------
    df : pandas.DataFrame
        The DataFrame to apply outlier capping to (can be training or test set)
    columns : list
        List of column names to apply outlier capping to
    bounds : dict
        Dictionary of pre-calculated bounds from training data
        
    Returns:
    --------
    df_capped : pandas.DataFrame
        A new DataFrame with outliers capped at the specified bounds
    """
    # Create a copy to avoid modifying the original DataFrame
    df_capped = df.copy()

    # Apply capping to each specified column
    for col in columns:
        # Extract pre-calculated bounds for this column
        lower_bound = ...
        upper_bound = ...
        
        # Use np.clip to cap values at the lower and upper bounds
        # - Values below lower_bound will be set to lower_bound
        # - Values above upper_bound will be set to upper_bound
        # - Values within the range remain unchanged
        df_capped[col] = ...

    return df_capped


# STEP 1: Define columns that need outlier handling
outlier_columns = ['chol', 'oldpeak']

# STEP 2: Calculate bounds using ONLY training data
outlier_bounds = ...

# STEP 3:
# Apply the bounds to training data
X_train_capped = ...

# Apply the SAME bounds to test data (prevents data leakage)
X_test_capped = ...

# IMPORTANT: Replace original dataframes with capped versions
X_train = ...
X_test = ...
```

```python
X_train.head()
```

**Why it's better to use one-hot encoding here?**

Even if categorical variables are converted to numbers (e.g., 0, 1, 2), ML models might misinterpret them as having mathematical meaning.

**Problems with Numeric Encoding**:
1. Ordinal vs. Nominal: Some categories have an order (ordinal), while others don’t (nominal). Treating them as numbers can be misleading.
2. False Mathematical Relationships: ML models may assume numerical significance (e.g., ‘thal’ = 3 is three times ‘thal’ = 1).
3. Distance Misinterpretation: Algorithms like SVM use distances, leading to incorrect relationships between categories.

```python
# One-Hot Encoding for Categorical Features
from sklearn.preprocessing import OneHotEncoder

# Step 1: Identify categorical features requiring one-hot encoding
encoded_features = ['cp', 'restecg', 'slope', 'thal', 'ca', 'age_group']

# Step 2: Initialize one-hot encoder (drop first category to avoid multicollinearity)
# Set drop='first' removes one column per feature to avoid perfect multicollinearity
# Set sparse_output = False, which is equivalent to calling .toarray()
encoder = ...

# Step 3: Fit on training data and transform both train and test sets
# Create a DataFrame with proper column names for the encoded features 
# - encoder.fit_transform() learns categories from AND transforms training data
# - encoder.get_feature_names_out() provides clear column names for the encoded features 
X_train_encoded = ...

# Important: We only transform test data (not fit_transform) to prevent data leakage
X_test_encoded = ...

# Step 4: Identify remaining features (those not selected for one-hot encoding)
remaining_features = ...

# Step 5: Concatenate remaining features and one-hot encoded categorical features
X_train = ...
X_test = ...
```

```python
# Standardize Numerical Features
from sklearn.preprocessing import StandardScaler
scaler = ...
numeric_features = ['age', 'trestbps', 'chol', 'thalach', 'oldpeak', 'age_chol']

X_train[numeric_features] = ...
X_test[numeric_features] = ...
```

```python
grader.check("q4")
```

### Introduction to RFE (Recursive Feature Elimination) [Codes Provided]
**Recursive Feature Elimination (RFE)** is a feature selection technique that helps identify the most important features for a machine learning model. It works by:

1. **Training a model** on all features.
2. **Ranking features** based on their importance (using model coefficients or feature importance scores).
3. **Eliminating the least important feature(s)** and repeating the process until the desired number of features is reached.

**Why use RFE?**

- Reduces overfitting by removing irrelevant features.
- Improves model interpretability.
- Enhances computational efficiency.

It is commonly used with models like **Logistic Regression**, **Decision Trees**, and **SVM** to optimize performance. In our example, we will use RFE with Logistic regression. 

The selected features are:
```python
['exang', 'cp_1', 'cp_2', 'cp_3', 'slope_1', 'thal_2', 'thal_3', 'ca_1', 'ca_2', 'ca_3']
```

```python
# Feature Selection using Recursive Feature Elimination (RFE)
from sklearn.feature_selection import RFE
from sklearn.linear_model import LogisticRegression

model = LogisticRegression(max_iter=1000)
rfe = RFE(model, n_features_to_select=10)  # Adjust based on desired number of features
X_train_selected = rfe.fit_transform(X_train, y_train)
X_test_selected = rfe.transform(X_test)  # Apply the same transformation to test set

# Print selected feature indices
selected_features = X_train.columns[rfe.support_]
print(f"Selected Features: {selected_features.tolist()}")
```

# Step 5: Model Training & Selection
**Tasks:**
- Train and evaluate multiple classification models using `X_train_selected`, `X_test_selected`, `y_train`, and `y_test.`
- Perform hyperparameter tuning using **GridSearchCV** to optimize model performance.
- Select the best model based on **cross-validation score** and evaluate it on test data.

```python
from sklearn.linear_model import LogisticRegression
from sklearn.tree import DecisionTreeClassifier
from sklearn.svm import SVC
from sklearn.metrics import accuracy_score, precision_score, recall_score, f1_score, roc_auc_score, confusion_matrix, classification_report
from sklearn.model_selection import cross_val_score, GridSearchCV
```

```python
# Step 1: Initialize Models & Parameter Grids (codes provided)
models = {
    'Logistic Regression': LogisticRegression(max_iter = 1000, random_state=42),
    'Decision Tree': DecisionTreeClassifier(random_state=42),
    'SVM': SVC(probability=True, random_state=42)
} 

param_grids = {
    'Logistic Regression': {
        'C': [0.01, 0.1, 1, 10, 100],
        'solver': ['liblinear', 'saga']
    },
    'Decision Tree': {
        'max_depth': [None, 5, 10, 15, 20],
        'min_samples_split': [2, 5, 10],
        'min_samples_leaf': [1, 2, 4]
    },
    'SVM': {
        'C': [0.1, 1, 10, 100],
        'kernel': ['linear', 'rbf', 'poly'],
        'gamma': ['scale', 'auto', 0.1, 0.01]
    }
}
```

```python
from sklearn.metrics import accuracy_score, classification_report
# Step 2: Hyperparameter Tuning with Cross-validation
best_models = {}
cv_scores = {}  # Store cross-validation scores for each model

for name, model in models.items():
    print(f"\nPerforming Grid Search for {name}")

    grid_search = GridSearchCV(
        ...
        ...
        cv=5,
        scoring='accuracy',
        n_jobs=-1
    )

    # Fit on training data
    ...

    # Store the best estimator
    best_models[name] = ...
    # Store the best cross-validation score
    cv_scores[name] = ...

    print(f"Best parameters: {grid_search.best_params_}")
    print(f"Best cross-validation score: {cv_scores[name]:.4f}")

# Step 3: Select the best model based on CV score
best_model_name = max(cv_scores, key=cv_scores.get)  # Model with highest CV score
# Retrieve the best trained model object from best_models dictionary
best_model = ...

print(f"\nBest model based on cross-validation: {best_model_name}")
```

```python
# Step 4: Evaluate the Best Model on Test Data
y_pred = ...

test_accuracy = ...

print(f"\nTest Accuracy of {best_model_name}: {test_accuracy:.4f}")
print("\nClassification Report:")
# print classification report
...
```

```python
grader.check("q5")
```

```python
# Learning Curves to Check for Overfitting/Underfitting (codes provided)
from sklearn.model_selection import learning_curve

def plot_learning_curve(estimator, title, X, y, cv=5):
    train_sizes_array = np.linspace(0.1, 1.0, 10)  # Renamed to avoid overwriting
    
    train_sizes, train_scores, test_scores = learning_curve(
        estimator, X, y, cv=cv, train_sizes=train_sizes_array, scoring='accuracy'
    )
    
    # Compute mean and standard deviation for both training and validation scores
    train_mean = np.mean(train_scores, axis=1)
    train_std = np.std(train_scores, axis=1)
    test_mean = np.mean(test_scores, axis=1)
    test_std = np.std(test_scores, axis=1)
    
    # Plot learning curve
    plt.figure(figsize=(10, 6))
    plt.title(title)
    plt.xlabel("Training examples")
    plt.ylabel("Accuracy Score")
    plt.grid()

    # Plot shaded areas representing standard deviation
    plt.fill_between(train_sizes, train_mean - train_std, train_mean + train_std, alpha=0.1, color="r")
    plt.fill_between(train_sizes, test_mean - test_std, test_mean + test_std, alpha=0.1, color="g")

    # Plot mean scores
    plt.plot(train_sizes, train_mean, 'o-', color="r", label="Training score")
    plt.plot(train_sizes, test_mean, 'o-', color="g", label="Cross-validation score")

    plt.legend(loc="best")
    plt.show()  # Ensure the plot is displayed inside the function

# Plot learning curve for the best model
plot_learning_curve(best_model, f"Learning Curve for {best_model_name}", X_train_selected, y_train)

```

**How to Interpret the Results**:

- Overfitting: If there's a large gap between training score (red) and cross-validation score (green), with training score much higher
- Underfitting: If both training and cross-validation scores are low and relatively close
- Good fit: If both training and cross-validation scores are high and relatively close

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1 (points: None)
```
>>> assert df.duplicated().sum() == 0, 'There should be no duplicates left after dropping them.'
>>> assert df.shape[0] == 302 and df.shape[1], 'The shape of the dataset is not correct.'
```


Tests for q2 (points: None)
```
>>> assert len(bin_edges) == 21, f'Expected 21 bin edges, got {len(bin_edges)}'
>>> assert sum(counts) == len(df['age']), f'Expected {len(df['age'])}, but got {sum(counts)}'
```
```
>>> unique_values = set(df['target'].unique())
>>> assert unique_values.issubset({0, 1}), f'Unexpected unique values in target: {unique_values}'
>>> plot_counts = [p.get_height() for p in ax_count.patches]
>>> actual_counts = df['target'].value_counts().sort_index().tolist()
>>> assert np.array_equal(plot_counts, actual_counts), f'Mismatch: Expected {actual_counts}, got {plot_counts}'
>>> x_labels = [t.get_text() for t in ax_count.get_xticklabels()]
>>> expected_labels = sorted(df['target'].unique().tolist())
>>> assert [int(label) for label in x_labels] == expected_labels, f'Expected {expected_labels}, got {x_labels}'
>>> assert sum(plot_counts) == len(df['target']), f'Expected {len(df['target'])}, but got {sum(plot_counts)}'
```
```
>>> corr_matrix = df.corr()
>>> heatmap_values = np.array([float(t.get_text()) for t in ax.texts]).reshape(corr_matrix.shape)
>>> assert np.allclose(heatmap_values, corr_matrix.round(1)), f'Mismatch between heatmap annotations and computed correlations: {heatmap_values}'
```
```
>>> chol_values = [line.get_xdata()[0] for line in ax1.lines[:3]]
>>> oldpeak_values = [line.get_xdata()[0] for line in ax2.lines[:3]]
>>> chol_stats = df['chol'].describe()
>>> oldpeak_stats = df['oldpeak'].describe()
>>> assert np.allclose(chol_values, [chol_stats['25%'], chol_stats['75%'], chol_stats['min']]), "Mismatch in boxplot statistics for 'chol'"
>>> assert np.allclose(oldpeak_values, [oldpeak_stats['25%'], oldpeak_stats['75%'], oldpeak_stats['min']]), "Mismatch in boxplot statistics for 'oldpeak'"
```


Tests for q3 (points: None)
```
>>> expected_results = {44: 'young', 45: 'young', 50: 'middle', 55: 'middle', 56: 'senior'}
>>> for age, expected_label in expected_results.items():
...     actual_label = df[df['age'] == age]['age_group'].values[0]
...     assert actual_label == expected_label, f'Test failed for age {age}: Expected {expected_label}, Got {actual_label}'
>>> assert df['age_chol'][:5].tolist() == [14679, 9250, 8364, 13216, 20178], 'The calculation of age_chol column is wrong.'
```


Tests for q4 (points: None)
```
>>> assert 'target' not in X.columns, "X should not contain the 'target' column"
>>> assert set(X.columns) == set(df.columns) - {'target'}, "X should contain all columns from df except 'target'"
>>> assert y.name == 'target', "y should be the 'target' column"
```
```
>>> sample_data = pd.DataFrame({'A': [10, 12, 15, 14, 200, 13, 12, 11, 300, 14], 'B': [100, 110, 120, 130, 125, 105, 115, 140, 150, 135]})
>>> bounds = calculate_outlier_bounds(sample_data, ['A', 'B'])
>>> lower_A, upper_A = (bounds['A']['lower'], bounds['A']['upper'])
>>> lower_B, upper_B = (bounds['B']['lower'], bounds['B']['upper'])
>>> Q1_A, Q3_A = (sample_data['A'].quantile(0.25), sample_data['A'].quantile(0.75))
>>> IQR_A = Q3_A - Q1_A
>>> expected_lower_A = Q1_A - 1.5 * IQR_A
>>> expected_upper_A = Q3_A + 1.5 * IQR_A
>>> Q1_B, Q3_B = (sample_data['B'].quantile(0.25), sample_data['B'].quantile(0.75))
>>> IQR_B = Q3_B - Q1_B
>>> expected_lower_B = Q1_B - 1.5 * IQR_B
>>> expected_upper_B = Q3_B + 1.5 * IQR_B
>>> assert lower_A == expected_lower_A and upper_A == expected_upper_A and (lower_B == expected_lower_B) and (upper_B == expected_upper_B), '`calculate_outlier_bounds` function does not work properly.'
>>> bounds = calculate_outlier_bounds(sample_data, ['A', 'B'])
>>> capped_df = apply_outlier_bounds(sample_data, ['A', 'B'], bounds)
>>> lower_A, upper_A = (bounds['A']['lower'], bounds['A']['upper'])
>>> assert all(lower_A <= capped_df['A']) and all(capped_df['A'] <= upper_A), 'Outliers were not capped correctly'
>>> unchanged_values = sample_data['A'].between(lower_A, upper_A)
>>> assert (sample_data['A'][unchanged_values] == capped_df['A'][unchanged_values]).all(), 'Non-outlier values should remain unchanged'
>>> assert capped_df['A'].max() <= upper_A and capped_df['A'].min() >= lower_A, 'Bounds not applied correctly'
>>> assert X_train_capped['chol'].max() == 366 and X_test_capped['chol'].max() == 366, 'Cholesterol (chol) should be capped at 366 for X_train and X_test after removing outliers. If your capped value is 376, please verify that you have correctly followed the train-test split instructions in the previous step.'
>>> assert X_train_capped['oldpeak'].max() == 4.5 and X_test_capped['oldpeak'].max() == 3.6, 'oldpeal should be capped at 4.5 for X_train and X_test after removing outliers.'
```
```
>>> for col in encoded_features:
...     assert col not in X_train.columns, f"Original categorical column '{col}' still exists in training set"
...     assert col not in X_test.columns, f"Original categorical column '{col}' still exists in test set"
>>> for col in remaining_features:
...     assert col in X_train.columns, f"Numerical column '{col}' is missing from training set"
...     assert col in X_test.columns, f"Numerical column '{col}' is missing from test set"
>>> encoded_cols = ['cp_1', 'cp_2', 'cp_3', 'restecg_1', 'restecg_2', 'slope_1', 'slope_2', 'thal_1', 'thal_2', 'thal_3', 'ca_1', 'ca_2', 'ca_3', 'ca_4', 'age_group_senior', 'age_group_young']
>>> train_encoded_values = X_train[encoded_cols].values.flatten()
>>> test_encoded_values = X_test[encoded_cols].values.flatten()
>>> assert set(np.unique(train_encoded_values)) <= {0, 1}, 'Training set encoded values are not binary'
>>> assert set(np.unique(test_encoded_values)) <= {0, 1}, 'Test set encoded values are not binary'
>>> assert X_train.shape[1] == 25, 'The shape of X_train after encoding is not correct'
>>> assert X_test.shape[1] == 25, 'The shape of X_test after encoding is not correct'
```
```
>>> assert np.allclose(X_train[numeric_features].mean(), 0, atol=1e-07), 'Mean is not close to 0 in training set'
>>> assert np.allclose(X_train[numeric_features].std(ddof=0), 1, atol=1e-07), 'Standard deviation is not close to 1 in training set'
>>> assert not np.allclose(X_test[numeric_features].mean(), 0, atol=1e-07), 'Test set should not have mean 0 after transformation'
```


Tests for q5 (points: None)
```
>>> expected_best_models_str = {'Logistic Regression': "LogisticRegression(C=10, max_iter=1000, random_state=42, solver='liblinear')", 'Decision Tree': 'DecisionTreeClassifier(max_depth=5, min_samples_split=10, random_state=42)', 'SVM': "SVC(C=1, kernel='linear', probability=True, random_state=42)"}
>>> expected_cv_scores = {'Logistic Regression': round(0.8507653061224489, 4), 'Decision Tree': round(0.7930272108843538, 4), 'SVM': round(0.8549319727891156, 4)}
>>> for key in expected_best_models_str:
...     assert str(best_models[key]) == expected_best_models_str[key], f'Mismatch in best_models for {key}'
>>> for key in expected_cv_scores:
...     assert round(cv_scores[key], 4) == expected_cv_scores[key], f'Mismatch in cv_scores for {key}'
```
```
>>> assert round(test_accuracy, 4) == 0.7869, 'Test accuracy is not correct.'
```
